# Classical probability distributions

**Exercise 1** (1 point). Recall that in classical con flipping, get heads with probability $P(X=0) = p_0$ and tails with $P(X=1) = p_1$ for each toss of the coin, where $p_i\geq 0$ for all $i$, and the probabilities sum to one: $\sum_i p_i = 1$. Create a sample with a 1000 data points using numpy, with a probability of getting tails being 0.3. This is the parameter that the `binomial` function takes. Store the outcome in an array called `x_data`.

```python
from assignment_helper import *

```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-21-7d7c8635f425> in <module>
----> 1 from assignment_helper import *

ModuleNotFoundError: No module named 'assignment_helper'
```

```python
n_samples = 1000
#

import numpy as np
p_1= 0.3
x_data= np.random.binomial(1, p_1, n_samples)

```

<!-- grade_id: score1, points: 1, grade: true, solution: false -->
```python
assert type(x_data) is np.ndarray
assert abs(p_1-x_data.sum()/n_samples) < 0.05

```

**Exercise 2** (1 point). As you recall, we may also write the probability distribution as a stochastic vector $\vec{p} = \begin{bmatrix} p_0 \\ p_1 \end{bmatrix}$. The normalization constraint on the probability distribution says that the norm of the vector is restricted to one in the $l_1$ norm. In other words, $||\vec{p}||_1 = \sum_i |p_i| = 1$. This would be the unit circle in the $l_1$ norm, but since $p_i\geq 0$, we are restricted to a quarter of the unit circle, just as we plotted above. Write a function that checks whether a given two-dimensional vector is a stochastic vector. That is, it should return `True` if all elements are positive and the 1-norm is approximately one, and it should return `False` otherwise. The input of the function is a numpy array.

```python
def is_stochastic_vector(p: np.array):
    #
   
    if p.sum()== 1:
        return True
    return False


```

<!-- grade_id: score2, points: 1, grade: true, solution: false -->
```python
assert is_stochastic_vector(np.array([0.2, 0.3])) is False
assert is_stochastic_vector(np.array([-0.2, 0.7])) is False
assert is_stochastic_vector(np.array([0.2, 0.8])) is True
```

**Exercise 3** (1 point). The probability of heads is just the first element in the $\vec{p}$ and we can use a projection to extract it. For the first element of the stochastic vector, the projection is described by the matrix $\begin{bmatrix} 1 & 0\\0 & 0\end{bmatrix}$. Write a function that performs this projection on a two-element vector described by a numpy array. Your output after the projection is also a two-element vector.

```python
def project_to_first_basis_vector(p: np.array):
    #
 
    stoch= np.array([[1, 0], [0, 0]])
    return p.dot(stoch)
```

<!-- grade_id: score3, points: 1, grade: true, solution: false -->
```python
assert np.alltrue(project_to_first_basis_vector(np.array([0.2, 0.3])) == np.array([0.2, 0.])) == True
assert np.alltrue(project_to_first_basis_vector(np.array([1., 0.])) == np.array([1., 0.])) == True
```

**Exercise 4** (1 point). The projection operators introduce some linear algebra to working with probability distributions. We can also use linear algebra to transform one probability distribution to another. A left *stochastic matrix* will map stochastic vectors to stochastic vectors when multiplied from the left: its columns add up to one. Write a function that takes a matrix and a vector as input arguments (both are numpy arrays), checks whether the vector is a stochastic vector and whether the matrix is left stochastic. If they are, return the matrix applied to the vector, otherwise raise a `ValueError`. You can call the function `is_stochastic_vector` that you defined above.

```python
def apply_stochastic_matrix(p: np.array, M: np.array):
    """Apply the matrix M to the vector p, but only if
    p is a stochastic vector and M is a left stochastic
    matrix. Otherwise raise a ValueError.
    """
    #
    
    if is_stochastic_vector(p)== True and matrix.sum(axis= 0).sum()== 2:
        return matrix
    raise ValueError
               

```

<!-- grade_id: score4, points: 1, grade: true, solution: false -->
```python
p = np.array([[.5], [.5]])
M = np.array([[0.7, 0.6], [0.3, 0.4]])
assert abs(np.linalg.norm(apply_stochastic_matrix(p, M), ord=1)-1) < 0.01
M = np.array([[0.7, 0.6], [0.3, 0.5]])
try:
    apply_stochastic_matrix(p, M)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

Output:
```
<ipython-input-22-4401ad3800c4> in <module>
      1 p = np.array([[.5], [.5]])
      2 M = np.array([[0.7, 0.6], [0.3, 0.4]])
----> 3 assert abs(np.linalg.norm(apply_stochastic_matrix(p, M), ord=1)-1) < 0.01
      4 M = np.array([[0.7, 0.6], [0.3, 0.5]])
      5 try:

<ipython-input-19-91fc74e74668> in apply_stochastic_matrix(p, M)
      6     #
      7 
----> 8     if is_stochastic_vector(p)== True and matrix.sum(axis= 0).sum()== 2:
      9         return matrix
     10     raise ValueError

NameError: name 'matrix' is not defined
```

**Exercise 5** (1 point). Create a left stochastic matrix in a variable called `M` that transforms the uniform distribution $\vec{p}= \begin{bmatrix} 0.5 \\ 0.5 \end{bmatrix}$ to $\begin{bmatrix} 0.6 \\ 0.4 \end{bmatrix}$. `M` should be a two-dimensional numpy array.

```python
#

M = np.array([[0.7, 0.5], [0.3, 0.5]])

```

<!-- grade_id: score5, points: 1, grade: true, solution: false -->
```python
assert np.alltrue(M.dot(np.array([0.5, 0.5])) == np.array([0.6, 0.4])) == True
```

**Exercise 6** (1 point). Calculate the entropy of this distribution $\begin{bmatrix} 0.6 \\ 0.4 \end{bmatrix}$ in a variable called `H`.

```python
#

D= np.array([0.6, 0.4])
H= 0
H= np.sum(-D*np.log(D))
H/np.log(2)

```

Output:
```
0.9709505944546688
```

# Quantum states

**Exercise 7** (1 point). A quantum state is a probability distribution. A qubit state is a distribution over two values, similar to the coin flipping in the classical state. A major difference is that the entries are complex numbers and the normalization is in the $l_2$ norm. Create a function similar to `is_stochastic_vector` that checks whether a vector is a valid quantum state. The input is a numpy array and the output should be boolean.

```python
def is_quantum_state(psi: np.array):
    #
    
        

```

Output:
```
File "<ipython-input-24-0c375521f8dc>", line 4
    
    ^
SyntaxError: unexpected EOF while parsing
```

<!-- grade_id: score7, points: 1, grade: true, solution: false -->
```python
assert is_quantum_state(np.array([1/np.sqrt(2), 1/np.sqrt(2)]))
assert is_quantum_state(np.array([-1/np.sqrt(2), 1/np.sqrt(2)]))
assert is_quantum_state(np.array([-1/3, 2*np.sqrt(2)/3]))
assert is_quantum_state(np.array([0.2, 0.8])) is False
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-25-77c9bcc0370f> in <module>
----> 1 assert is_quantum_state(np.array([1/np.sqrt(2), 1/np.sqrt(2)]))
      2 assert is_quantum_state(np.array([-1/np.sqrt(2), 1/np.sqrt(2)]))
      3 assert is_quantum_state(np.array([-1/3, 2*np.sqrt(2)/3]))
      4 assert is_quantum_state(np.array([0.2, 0.8])) is False

NameError: name 'is_quantum_state' is not defined
```

**Exercise 8** (1 point). While working with numpy arrays is convenient, it is better to use a framework designed for quantum computing, since it often allows us to execute a circuit directly on a quantum computer. In your preferred framework, implement a circuit of a single qubit with no operation on it. You should create it in an object called `circuit`. Do not add a measurement. The evaluation will automatically branch according to which framework you chose.

```python
#

```

<!-- grade_id: score8, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert abs(amplitudes[0]-1.0) < 0.01
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-27-c50e1f34bcb8> in <module>
----> 1 amplitudes = get_amplitudes(circuit)
      2 assert abs(amplitudes[0]-1.0) < 0.01

NameError: name 'get_amplitudes' is not defined
```

**Exercise 9** (1 point). In the execution branching above, you see that we use the wavefunction simulator. This allows us to use the probability amplitudes as usual numpy arrays, as you can see above. If we ran the circuit on an actual quantum device, we would not be able to inspect the wavefunction, but we would have to rely on the statistics of measurements to understand what is happening in the circuit.

Create a circuit in your preferred framework that creates an equal superposition in a qubit using a Hadamard gate. Again, the name of the object should be `circuit`. The evaluation will be based on measurement statistics. In this case, you should explicitly specify the measurement on the qubit

```python
#
# YOUR CODE HERE

```

<!-- grade_id: score9, points: 1, grade: true, solution: false -->
```python
counts = get_counts(circuit)
assert abs(counts['0']/100-.5) < 0.2
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-29-6a1897bd8e9b> in <module>
----> 1 counts = get_counts(circuit)
      2 assert abs(counts['0']/100-.5) < 0.2

NameError: name 'get_counts' is not defined
```

**Exercise 10** (1 point). If you plotted the state before measurement on the Bloch sphere, it would have been on the equator halfway between the $|0\rangle$ and $|1\rangle$ states, and the tip of the X axis. If you apply the Hadamard on the $|1\rangle$, it would have been the point on the opposite and of the X axis, since the resulting superposition would have had a -1 amplitude for $|1\rangle$. The measurement statistics, however, would be identical. The negative sign plays a role in interference: for instance, applying a Hadamard again, would take you back to $|1\rangle$. Create the superposition after applying the Hadamard gate on $|1\rangle$. We will verify whether it picked up the phase. Do not include a measurement, since we will inspect the wavefunction.

```python
#
# YOUR CODE HERE


```

<!-- grade_id: score10, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert abs(amplitudes[1]+np.sqrt(2)/2) < 0.01
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-31-fde6ac009c5a> in <module>
----> 1 amplitudes = get_amplitudes(circuit)
      2 assert abs(amplitudes[1]+np.sqrt(2)/2) < 0.01

NameError: name 'get_amplitudes' is not defined
```

# More qubits and entanglement

**Exercise 11** (1 point). To get a sense of multiqubit states, it is important to be confident with the tensor product operation. Create a function that returns the four basis vectors, $|00\rangle$, $|01\rangle$, $|10\rangle$, and $|11\rangle$, of the tensor product space $\mathbb{C}^2\otimes\mathbb{C}^2$. The order in which they appear does not matter. The return value should be a list of four numpy arrays.

```python
def create_canonical_basis():
    #
    
    zer_ket= np.array([1, 0])
    one_ket= np.array([0, 1])
    zerozero_ket= np.concatenate((zero_ket[0]* zero_ket, zero_ket[1]* zero-ket))
    zeroone_ket= np.concatenate((zero_ket[0]* one_ket, zero_ket[1]* one-ket))
    onezero_ket= np.concatenate((one_ket[0]* zero_ket, one_ket[1]* zero-ket))
    oneone_ket= np.concatenate((one_ket[0]* oneo_ket, one_ket[1]* one-ket))        
```

```python
print(create_canonical_basis())
```

Output:
```
<ipython-input-35-7122ca5cf018> in <module>
----> 1 print(create_canonical_basis())

<ipython-input-34-78274201be04> in create_canonical_basis()
      4     zer_ket= np.array([1, 0])
      5     one_ket= np.array([0, 1])
----> 6     zerozero_ket= np.concatenate((zero_ket[0]* zero_ket, zero_ket[1]* zero-ket))
      7     zeroone_ket= np.concatenate((zero_ket[0]* one_ket, zero_ket[1]* one-ket))
      8     onezero_ket= np.concatenate((one_ket[0]* zero_ket, one_ket[1]* zero-ket))

NameError: name 'zero_ket' is not defined
```

**Exercise 12** (1 point). A generic product state has the form  $\begin{bmatrix}a_0b_0\\ a_0b_1\\ a_1b_0\\ a_1b_1\end{bmatrix}=a_0b_0|00\rangle + a_0b_1|01\rangle + a_1b_0|10\rangle + a_1b_1|11\rangle$ on $\mathbb{C}^2\otimes\mathbb{C}^2$, but not all. We can use the basis vectors to form vectors in the space that do not have a product structure. These are entangled states that show strong correlations. Entanglement is an important resource in quantum computing and being able to create a circuit that generates an entangled state is critical. Implement a circuit in your preferred framework to create the $|\phi^-\rangle = \frac{1}{\sqrt{2}}(|00\rangle-|11\rangle)$ state, that is, almost the same as the $|\phi^+\rangle$ state, but with the opposite sign of the probability amplitude of $|11\rangle$. Do not include a measurement, as we will verify the state with the wavefunction simulator

```python
statevctor_sim= BasicAer.get_backend("statevector_sim")
```

```python
#


bachend= statevector_sim
qr= QuantumRegister(2)
cr= ClassicalRegister(2)
qc= QuantumCircuit(qr,cr)


```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-36-b4bf0aa96d97> in <module>
      1 #
      2 
----> 3 statevctor_sim= BasicAer.get_backend("statevector_sim")
      4 bachend= statevector_sim
      5 qr= QuantumRegister(2)

NameError: name 'BasicAer' is not defined
```

<!-- grade_id: score12, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert all(np.isclose(np.array([np.sqrt(2)/2, 0, 0, -np.sqrt(2)/2]), amplitudes))
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-37-7f00a01f13bd> in <module>
----> 1 amplitudes = get_amplitudes(circuit)
      2 assert all(np.isclose(np.array([np.sqrt(2)/2, 0, 0, -np.sqrt(2)/2]), amplitudes))

NameError: name 'get_amplitudes' is not defined
```

# Measurements

**Exercise 1** (1 point). Measurements in the quantum computers we can access today always measure in the computational basis. This means that either the projection $|0\rangle\langle 0|$ or the $|1\rangle\langle 1|$ is applied on the qubit we are measuring, corresponding to the outcome 0 or 1. The Born rule tells us that we get either of these with the probability of the absolute value of the probability amplitde squared. The qubit afterwards is $|0\rangle$ or $|1\rangle$, respectively.

First, create circuit in your preferred framewrok with one quantum and two classical registers. The object should be called `circuit`.

```python
qasm_sim= BasicAer.get_backend("qasm_sim")
```

```python
#


backend= qasm_sim
qr= QuantumRegister(1)
cr1= ClassicalRegister(1)
cr2= ClassicalRegister(1)
qc= QuantumCircuit(qr, cr1, cr2)


```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-39-b9d4f89731cc> in <module>
      1 #
      2 
----> 3 qasm_sim= BasicAer.get_backend("qasm_sim")
      4 backend= qasm_sim
      5 qr= QuantumRegister(1)

NameError: name 'BasicAer' is not defined
```

**Exercise 2** (1 point). Next, extend the circuit with a Hadamard gate and a measurement on the qubit. Write the measurement result in the first classical register. You will get a probabilistic outcome of 0 or 1.

```python
#

qc= QuantumCircuit(qr, cr1, cr2)
qc.h(qr[0])
qc.measure(qr[0], cr1[0])
result= execute(qc, backend).result()
print(result.get_counts(qc))

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-41-53e81e0dc2c3> in <module>
      1 #
      2 
----> 3 qc= QuantumCircuit(qr, cr1, cr2)
      4 qc.h(qr[0])
      5 qc.measure(qr[0], cr1[0])

NameError: name 'QuantumCircuit' is not defined
```

**Exercise 3** (1 point). To see that the quantum state collapses to the basis state indicated by the outcome, apply a second measurement on the same circuit, but write the result in the second register. Executing the program should always give the same outcome in the two subsequent measurements.

```python
#

qc.measure(qr[0], cr2[0])
result = execute(qc, backend).result()
print(result.get_counts(qc))


```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-46-87d6bb6fc60d> in <module>
      1 #
      2 
----> 3 qc.measure(qr[0], cr2[0])
      4 result = execute(qc, backend).result()
      5 print(result.get_counts(qc))

NameError: name 'qc' is not defined
```

# Measuring multiqubit systems

**Exercise 4** (1 point). We typically work with local measurements, meaning that the qubits are measured separately. Create a two-qubit circuit with measurements included that reproduces the uniform distribution on four outcomes.

```python
#
# YOUR CODE HERE

```

**Exercise 5** (1 point). This is a typical product state, showing no correlation between the qubits. In contrast, if you make simultaneous measurements on an entangled state, for instance the $|\phi^+\rangle$ or the $|\phi^-\rangle$ state, the measurement outcomes will be correlated. Create either of these states and measure both qubits.

```python
#
# YOUR CODE HERE

```

Notice that we only observe 00 and 11 as outcomes, even though we made measurements on two spatially separated qubits. If one measurement gives a value, the other one always give the same. 

# Mixed states

**Exercise 6** (1 point). The density matrix formalism is critical in understanding noise models and decoherence, so it is indispensible to start working with current and near-future quantum computers. Many simulator backends provide noise models, but it is out of scope for this course to develop a deep understanding of it.

Instead, let's use the wavefunction simulator to create the (maximally) entangled state $|\phi^+\rangle$. The object should be called circuit and do not include measurements. Your circuit should be in an object called `circuit`.

```python
#

backend= statevector_sim
qr= QuantumRegister(2)
cr= ClassicalRegister(2)
qc= QuantumCircuit(qr, cr)
qc.h(qr[0])
qc.cx(qr[0], qr[1])
result = execute(circuit, backend).result()
amplitudes = result.get_statevector(qc)
print(result.get_statevector(qc))

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-49-57954e0443a4> in <module>
      1 #
      2 
----> 3 backend= statevector_sim
      4 qr= QuantumRegister(2)
      5 cr= ClassicalRegister(2)

NameError: name 'statevector_sim' is not defined
```

**Exercise 7** (1 point). Create the density matrix called `dm` of the probability amplitudes contained in the `amplitudes` array. Recall that you need a ket and a bra in this order to produce it, but the shape of the `amplitudes` array is incorrect for transposition. Reshape it first

```python
#

reshape= amplitudes.reshape(2, 2)
dm= reshape.dot(reshaped.T.conj())
dm

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-50-edcd73014669> in <module>
      1 #
      2 
----> 3 reshape= amplitudes.reshape(2, 2)
      4 dm= reshape.dot(reshaped.T.conj())
      5 dm

NameError: name 'amplitudes' is not defined
```

Taking the partial trace of a density matrix is the equivalent of taking the marginal of a joint probability distribution along one random variable. If we take the partial trace of this density matrix in any of the qubit subsystems, we are going to get the maximally mixed state, that is, the uniform distribution over the remaining system. This reflects the strong correlation between the two subsystems: if we marginalize over one, we have no predictive power over the other one. We exploit this property when we create protocols for preparing thermal states.


# Unitary evolution

**Exercise 1** (1 point). The conjugate transpose of a unitary matrix is its inverse. Write a function that checks if a $2\times 2$ numpy array is unitary.

```python
def is_unitary(matrix):
    
    #
   
    return np.array_equal(np.round(matrix.T.conj(), 8), np.round(np.linalg.inv(matrix), 8))

```

```python
X = np.array([[0, 1], [1, 0]])
H = 1/np.sqrt(2)*np.array([[1, 1], [1, -1]])
random_matrix = np.random.rand(2, 2)
assert is_unitary(X)
assert is_unitary(H)
assert is_unitary(random_matrix) is False
```

**Exercise 2** (1 point). The eigenvalues of unitary operators always lie on the complex unit circle. We exploit this property over and over again. For instance, the quantum phase estimation algorithm estimates the eigenvalues through the phase that describes the eigenvalues angle on the complex unit circle. Calculate the eigenvalues of the Hadamard gate and store them in a numpy array called `eigenvalues`.

```python
#

H= np.array([[1, 1], [1, -1]]/np.sqrt(2))
eval= np.linalg.eval(H)
eval
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-53-622f04246b94> in <module>
      2 
      3 H= np.array([[1, 1], [1, -1]]/np.sqrt(2))
----> 4 eval= np.linalg.eval(H)
      5 eval

AttributeError: module 'numpy.linalg' has no attribute 'eval'
```

**Exercise 3** (1 point). Unitary operations are reversible. In fact, we often reverse them on purpose: this is called uncomputing, and used, for instance, in the quantum matrix inversion algorithm that underlies many coherent quantum machine learning protocols. The reason we do that is to remove the entanglement we might have created, so that tracing out a subsystem would not interfere with our calculations (the remaining system will not be a mixed state). Create a circuit in your preferred framework that applies the X and H gates in this sequence. The circuit should be in an object called `circuit`.

```python
#

backend = statevector_sim
qr= QuantumRegister(1)
cr= ClassicalRegister(1)
qc= QuantumCircuit(qr, cr)
qc.x(qr[0])
qc.h(qr[0])

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-54-da66704738df> in <module>
      1 #
      2 
----> 3 backend = statevector_sim
      4 qr= QuantumRegister(1)
      5 cr= ClassicalRegister(1)

NameError: name 'statevector_sim' is not defined
```

```python
from assignment_helper import get_amplitudes
amplitudes = get_amplitudes(circuit)
assert all(np.isclose(amplitudes, np.array([1/np.sqrt(2), -1/np.sqrt(2)])))
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-55-0ec73fb326a5> in <module>
----> 1 from assignment_helper import get_amplitudes
      2 amplitudes = get_amplitudes(circuit)
      3 assert all(np.isclose(amplitudes, np.array([1/np.sqrt(2), -1/np.sqrt(2)])))

ModuleNotFoundError: No module named 'assignment_helper'
```

**Exercise 4** (1 point). Extend this circuit to reverse the calculations you have done.

```python
#

qc.h(qc[0])
qc.x(qc[0])

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-56-c78c63410b9e> in <module>
      1 #
      2 
----> 3 qc.h(qc[0])
      4 qc.x(qc[0])

NameError: name 'qc' is not defined
```

```python
amplitudes = get_amplitudes(circuit)
program_length = get_circuit_length(circuit)
assert all(np.isclose(amplitudes, np.array([1, 0]))), "The circuit is not successfully reversed"
assert program_length == 4, "The circuit does not have the right number of gates"
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-57-2d63332abf34> in <module>
----> 1 amplitudes = get_amplitudes(circuit)
      2 program_length = get_circuit_length(circuit)
      3 assert all(np.isclose(amplitudes, np.array([1, 0]))), "The circuit is not successfully reversed"
      4 assert program_length == 4, "The circuit does not have the right number of gates"

NameError: name 'get_amplitudes' is not defined
```